---
title: 'Знакомство с tidyverse'
description: 'В этой главе вы познакомитесь с основами tidyverse, научитесь загружать данные и проводить их первичный анализ.'
---

## 1. Загружаем пакет

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}
# Загрузите gfr
```

`@solution`
```{r}
library(tidyverse)
```

`@sct`
```{r}

```
